```python
# Initialize Otter
import otter
grader = otter.Notebook("check_exploring.ipynb")
```

# Instructions

This assignment lets you both define a strategy for picking the next point to explore and determine how you want to chop up a full path into way points. You'll need **path_planning.py** and **exploring.py** as well (for calculating the paths)

Note that there isn't a "right" answer for any of these. This is (mostly) a light-weight way to check your code for obvious problems before trying it in ROS. It's set up to make it easy to download a map and try some robot starting/ending points

Given to you:
- Image handling
- plotting
- Some structure for keeping/changing waypoints and converting to/from the map to the robot's coordinate space

Slides: https://docs.google.com/presentation/d/1tOUGoNDbzQeuzRwG-kJ-ipqDsCy3ABkIvNEdFmozRT8/edit?usp=sharing


```python
# The ever-present numpy and matplot lib
import numpy as np
import os
```

```python
# These commands will force JN to actually re-load the external file when you re-execute the import command
%load_ext autoreload
%autoreload 2
```

# Run exploring a SLAM map



```python
try:
    from path_planning import dijkstra, plot_with_path, open_image
    from exploring import find_all_possible_goals, find_best_point, plot_with_explore_points, find_waypoints, test_unseen, test_best
except:
    cur_dir = os.getcwd()
    if "Assignments" in cur_dir or "Skills" in cur_dir:
        path_planning_dir = "../ros_ws/src/lab3/lab3"

        os.chdir(path_planning_dir)
        from path_planning import dijkstra, plot_with_path, open_image
        from exploring import find_all_possible_goals, find_best_point, plot_with_explore_points, find_waypoints, test_unseen, test_best
        os.chdir(cur_dir)

```

```python
# You can use map or initial
_, im_thresh = open_image("map.pgm")

# Split into two pieces - find all of the points that are adjacent to an unknown area and a visible area. The
#. next question will pick the "best" one
all_unseen = find_all_possible_goals(im_thresh)
assert test_unseen(im_thresh, all_unseen)

# Note: This can be a little slow (it's drawing a lot of little points)
zoom = 1.0
plot_with_explore_points(im_thresh, zoom=1.0, robot_loc=None, explore_points=all_unseen, best_pt=None)
```

Output:
```
c:\Users\russd\.vscode\ROB456 FINAL\ROB456\ros_ws\src\lab3\lab3
Cell In[5], line 7
      4 # Split into two pieces - find all of the points that are adjacent to an unknown area and a visible area. The
      5 #. next question will pick the "best" one
      6 all_unseen = find_all_possible_goals(im_thresh)
----> 7 assert test_unseen(im_thresh, all_unseen)
      9 # Note: This can be a little slow (it's drawing a lot of little points)
     10 zoom = 1.0

File c:\Users\russd\.vscode\ROB456 FINAL\ROB456\ros_ws\src\lab3\lab3\exploring.py:134, in test_unseen(im, pts)
    133 def test_unseen(im, pts):
--> 134     for pt in pts:
    135         count_free = 0
    136         count_unseen = 0

TypeError: 'NoneType' object is not iterable
```

```python
grader.check("unknown_area")
```

## Find best
No right naswer here, but going to test for a pixel that has at least 3 free neighbors

```python
# Where the robot is might influence your choice. Shows path to selected point.
robot_start_loc = (60, 40)
best_unseen = find_best_point(im_thresh, all_unseen, robot_loc=robot_start_loc)
assert test_best(im_thresh, best_unseen)

path = dijkstra(im_thresh, robot_start_loc, best_unseen)
plot_with_path(im_thresh, zoom=1.0, robot_loc=robot_start_loc, goal_loc=best_unseen, path=path)

```

```python
grader.check("find_best")
```

# Run waypoint generation on a pretend SLAM map
Note: This code assumes you have the previous problem completed

```python
path = dijkstra(im_thresh, robot_start_loc, best_unseen)
waypoints = find_waypoints(im_thresh, path)
plot_with_path(im_thresh, zoom=1.0, robot_loc=robot_start_loc, goal_loc=best_unseen, path=waypoints)
print(f"Waypoints has {len(waypoints)} points of {len(path)}")
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[6], line 1
----> 1 path = dijkstra(im_thresh, robot_start_loc, best_unseen)
      2 waypoints = find_waypoints(im_thresh, path)
      3 plot_with_path(im_thresh, zoom=1.0, robot_loc=robot_start_loc, goal_loc=best_unseen, path=waypoints)

NameError: name 'robot_start_loc' is not defined
```

```python
grader.check("Waypoints")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS W26 (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5
```

```python
grader.check("hours_collaborators")
```

### To submit

- Do a restart then run all to make sure everything runs ok
- Save the file
- Submit this .ipynb file and **path_planning.py** and **exploring.py** through gradescope, Lab 4: Explore
- You do NOT need to submit the data files - we will supply those
- Take out/suppress all print statements

If the Gradescope autograder fails, please check here first for common reasons for it to fail
    https://docs.google.com/presentation/d/1tYa5oycUiG4YhXUq5vHvPOpWJ4k_xUPp2rUNIL7Q9RI/edit?usp=sharing

Most likely failure for this assignment is not naming the data directory and files correctly; capitalization matters for the Gradescope grader. 

Second most likely failure is because your code ran very, very slowly and the autograder timed out.

Third reason for failure is you forgot to include **path_planning.py** and **exploring.py** and/or you put the file(s) in a folder

Tests for Waypoints (points: 3)
```
>>> assert len(waypoints) < len(path)
```


Tests for find_best (points: 2)
```
>>> assert test_best(im_thresh, best_unseen)
```


Tests for hours_collaborators (points: 1)
```
>>> assert not 'not filled out' in worked_with_names
```
```
>>> assert not 'not filled out' in websites
```
```
>>> assert hours > 0
```


Tests for unknown_area (points: 5)
```
>>> assert test_unseen(im_thresh, all_unseen)
```
